```python
# Examen 1: Puntaje de Concierto
Escribe una función en Python llamada `calcular_puntaje` que reciba el número de notas perfectas y el multiplicador de combo de un concierto de Togenashi Togeari, y devuelva el puntaje total (multiplicando ambos valores).
```

<!-- grade_id: cell-43ae3ee7373e3111, points: 0, grade: false, solution: true -->
```python
def calcular_puntaje(notas_perfectas, combo):
       ### BEGIN SOLUTION
       return notas_perfectas * combo
       ### END SOLUTION
```

<!-- grade_id: cell-7a2a48798877ceae, points: 10, grade: true, solution: false -->
```python
assert calcular_puntaje(100, 2) == 200
   assert calcular_puntaje(50, 5) == 250
   assert calcular_puntaje(0, 10) == 0
   print("¡Respuesta correcta, el concierto fue un éxito!")
```